```python
%matplotlib inline
```

```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt

from sklearn.linear_model import LinearRegression, RANSACRegressor, LogisticRegression
from sklearn.model_selection import train_test_split, GridSearchCV, RandomizedSearchCV
from sklearn.preprocessing import MinMaxScaler
from sklearn.metrics import classification_report, RocCurveDisplay, precision_score, recall_score, f1_score
# from nose.tools import *
# Write your imports in the cell below
```

```python
def assert_is_not_none(value, error_message="Value cannot be None"):
  if value is None:
    raise AssertionError(error_message)
  
def assert_equal(actual_shape, expected_shape):
  if expected_shape != actual_shape:
    message = f"Shapes are not equal: expected {expected_shape}, got {actual_shape}"
    raise AssertionError(message)
  
def assert_almost_equal(actual, expected, delta=0.05):
  message = f"Values are not almost equal: expected {expected}, got {actual} (difference: {abs(actual - expected):.6f})"
  assert abs(actual - expected) <= delta, message

def assert_greater(actual_value, threshold, message="Value must be greater than the threshold"):
  if actual_value <= threshold:
    raise AssertionError(message)
  
def assert_less(actual_value, threshold):
  if actual_value >= threshold:
    raise AssertionError("Value must be lower than the threshold")
```

```python
np.random.seed(1234)
```

# Model Training and Improvement Lab
## Comparing and selecting models

### 1. Read the data (1 point)
Like in the previous lab, you need to read the Portuguese bank dataset [here](https://archive.ics.uci.edu/ml/machine-learning-databases/00222/). It has been provided for you in the `data` folder.

Read the dataset using `pandas` (you can use the library with the alias `pd`). Save it in the `bank_data` variable.

<!-- grade_id: cell-39fe3843121d7b36, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("./data/bank.csv", sep=";")
bank_data
```

Output:
```
age            job  marital  education default  balance housing loan  \
0      30     unemployed  married    primary      no     1787      no   no   
1      33       services  married  secondary      no     4789     yes  yes   
2      35     management   single   tertiary      no     1350     yes   no   
3      30     management  married   tertiary      no     1476     yes  yes   
4      59    blue-collar  married  secondary      no        0     yes   no   
...   ...            ...      ...        ...     ...      ...     ...  ...   
4516   33       services  married  secondary      no     -333     yes   no   
4517   57  self-employed  married   tertiary     yes    -3313     yes  yes   
4518   57     technician  married  secondary      no      295      no   no   
4519   28    blue-collar  married  secondary      no     1137      no   no   
4520   44   entrepreneur   single   tertiary      no     1136     yes  yes   

       contact  day month  duration  campaign  pdays  previous poutco
```

<!-- grade_id: cell-b183f83b98f16a29, points: 1, grade: true, solution: false -->
```python
# From now on, all test cells might contain hidden tests. If you follow the instructions correctly, 
# your solution will be graded with maximum points
assert_is_not_none(bank_data)
```

### 2. Preprocess the data (1 point)
Separate explanatory features from labels. Save all features (16 columns total) in the variable `bank_features`. Save the labels (corresponding to the `y` column) in the `bank_labels` variable. Rewrite the labels to be `0` and `1` instead of `no` and `yes`: `bank_labels` should be a numeric column.

<!-- grade_id: cell-fb1c7a38947afc06, points: 0, grade: false, solution: true -->
```python
bank_features = bank_data.drop(columns="y")
bank_labels = bank_data.y
bank_labels = np.where(bank_labels == "yes", 1, 0)
```

<!-- grade_id: cell-30a215c478663fb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features)
assert_is_not_none(bank_labels)
```

### 3. Get indicator variables (1 point)
Get indicator (dummy) variables for all categorical columns in `bank_features`. Overwrite the `bank_features` variable to store the new data.

<!-- grade_id: cell-8724b6857c2eb7fc, points: 0, grade: false, solution: true -->
```python
bank_features = pd.get_dummies(bank_features)
bank_features.shape
```

Output:
```
(4521, 51)
```

<!-- grade_id: cell-c2b2c0bd2d30ea5e, points: 1, grade: true, solution: false -->
```python
assert_equal(bank_features.shape, (4521, 51))
```

### 4. Split the data (1 point)
Split the data into training and testing set, with 70% of the data for training. Because the output labels are not equaly distributed, use stratification based on the `bank_labels`.

<!-- grade_id: cell-f46c32bb67654092, points: 0, grade: false, solution: true -->
```python
bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = train_test_split(
    bank_features,
    bank_labels,
    train_size=0.7, 
    test_size=0.3, 
    random_state=42, 
    stratify=bank_labels
)
bank_features_train.shape, bank_labels_train.shape, bank_features_test.shape, bank_labels_test.shape
```

Output:
```
((3164, 51), (3164,), (1357, 51), (1357,))
```

<!-- grade_id: cell-b14f41fb226c0f37, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features_train)
assert_is_not_none(bank_labels_train)
assert_is_not_none(bank_features_test)
assert_is_not_none(bank_labels_test)
```

### 5. Train a baseline algorithm (1 point)
Train a logistic regression using the training data. Use 1 000 000 (`1e6`) as the value of C. Score it using the testing data. Save the score in the `baseline_score` variable. You should see a fairly high score.

<!-- grade_id: cell-2fd7633e6266fcba, points: 0, grade: false, solution: true -->
```python
model = LogisticRegression(C = 1e6)
model.fit(bank_features_train, bank_labels_train)
baseline_score = model.score(bank_features_test, bank_labels_test)


baseline_score
```

Output:
```
c:\Boyko Boev\SoftUni-Machine-Learning-September-2022\.venv\Lib\site-packages\sklearn\linear_model\_logistic.py:469: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
0.89240972733972
```

<!-- grade_id: cell-a3af8c95790cccc4, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
assert_greater(baseline_score, 0.7)
```

### 6. Select a better score (2 points)
As you alrady saw, the positive examples are very few. If you aren't convinced, just check the counts.

We know that the default scoring (accuracy) isn't correct in this case. Better measures would be precision and recall. However, we only want one number. Evaluate the algorithm once again, using a standard scoring method which combines precision and recall. Overwrite the `baseline_score` variable.

Don't forget to score the model on the testing data only.

<!-- grade_id: cell-9b44f24da39ba641, points: 0, grade: false, solution: true -->
```python

bank_labels_predicted = model.predict(bank_features_test)
print(classification_report(bank_labels_test, bank_labels_predicted))

baseline_score = f1_score(bank_labels_predicted, bank_labels_test)
print(baseline_score)
```

Output:
```
precision    recall  f1-score   support

           0       0.91      0.98      0.94      1201
           1       0.59      0.21      0.31       156

    accuracy                           0.89      1357
   macro avg       0.75      0.60      0.63      1357
weighted avg       0.87      0.89      0.87      1357

0.3113207547169811
```

<!-- grade_id: cell-9f4aca08f036d9de, points: 2, grade: true, solution: false -->
```python
assert_less(baseline_score, 0.7)
```

### 7. Tune your model (2 points)
Fine-tune the `C` and `max_iter` parameters.

Use full grid search with the following values:
* `C`: 0.0001, 0.01, 0.1, 1, 10, 100, 10000
* `max_iter`: 50, 100, 300, 1000
* `fit_itercept`: True, False

Save the grid search result in the `grid_search` variable. Don't forget to use the better scoring model that you obtained in the previous task.

<!-- grade_id: cell-88cb1341dfd72476, points: 0, grade: false, solution: true -->
```python
grid_search = GridSearchCV(
    estimator=LogisticRegression(),
    param_grid={
     "C": [0.0001, 0.01, 0.1, 1, 10, 100, 10000],
     "max_iter": [50, 100, 300, 1000],   
     "fit_intercept": [True, False]
    }
)
grid_search.fit(bank_features_test, bank_labels_test)
```

Output:
```
c:\Boyko Boev\SoftUni-Machine-Learning-September-2022\.venv\Lib\site-packages\sklearn\linear_model\_logistic.py:469: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-learn.org/stable/modules/linear_model.html#logistic-regression
  n_iter_i = _check_optimize_result(
c:\Boyko Boev\SoftUni-Machine-Learning-September-2022\.venv\Lib\site-packages\sklearn\linear_model\_logistic.py:469: ConvergenceWarning: lbfgs failed to converge (status=1):
STOP: TOTAL NO. of ITERATIONS REACHED LIMIT.

Increase the number of iterations (max_iter) or scale the data as shown in:
    https://scikit-learn.org/stable/modules/preprocessing.html
Please also refer to the documentation for alternative solver options:
    https://scikit-le
```

<!-- grade_id: cell-0e7311dbe13cf894, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(grid_search.best_estimator_)
```

### 8. Compare scores (1 point)
Use the best estimator from your grid search. Score it using the function from problem 6. Save your answer in `tuned_score`.

<!-- grade_id: cell-98568f584f6c0778, points: 0, grade: false, solution: true -->
```python
print(
    # grid_search.best_estimator_.score(),
    grid_search.best_estimator_,
    grid_search.best_estimator_.fit_intercept,
    grid_search.best_estimator_.max_iter,
)

tuned_predictions = grid_search.predict(bank_features_test)
tuned_score = f1_score(tuned_predictions, bank_labels_test)
```

Output:
```
LogisticRegression(C=100, fit_intercept=False, max_iter=1000) False 1000
```

<!-- grade_id: cell-545b4d41fef4eeef, points: 1, grade: true, solution: false -->
```python
print(tuned_score)
```

Output:
```
0.3394495412844037
```

```python
print(baseline_score - tuned_score)
```

Output:
```
-0.02812878656742257
```

Hmmmm, it seems we have not obtained a better algorithm, even the opposite (the difference is marginal and depends on the random initialization of the cross-validation datasets).

We can, of course, do a lot more things to improve our model's performance, such as normalizing the data, feature selection and feature engineering, trying out different aspects, e.g. polynomial terms, RANSAC; even boosting (we'll talk about this later). However, we'll stop at this point.

What can we conclude? It seems that this is close to the best performance we can get out of this algorithm, given these data points.

We can try improving (cleaning) our dataset, selecting features, etc. but we most likely need a better algorithm. In the next labs, we're going to explore that.